---
title: "Analyse factorielle multiple (AFM)"
author: "Guyliann Engels & Philippe Grosjean"
description: "**SDD II Module 8** : Analyse factorielle multiple (AFM)."
tutorial:
  id: "B08Lb_mfa"
  version: 2.0.1/8
output: 
  learnr::tutorial:
    progressive: true
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
BioDataScience2::learnr_setup()
# Package utile
SciViews::R()
# Functions utiles 
# Code mfa() version 2.0.0
# Copyright (c) 2020, Philippe Grosjean (phgrosjean@sciviews.org)

mfa <- function(data, formula, nd = 5, suprow = NA, ..., graph = FALSE) {
  # TODO: rework Call in the final MFA object
  if (is.na(suprow)) suprow <- NULL # MFA uses NULL instead of NA!
  if (!rlang::is_formula(formula))
    stop("'formula' must be a formula object")
  if (!is.null(rlang::f_lhs(formula)))
    stop("'formula cannot have left-hand side (must be ~ n1 * type...")
  params <- get_groups(list(expr = rlang::f_rhs(formula)))
  
  # data must be a data frame
  if (!inherits(data, "data.frame"))
    stop("'data' must be a data.frame")
  data <- as.data.frame(data) # No tibble or data.table!
  # Number of columns in  data must match specifications in the formula
  if (ncol(data) != sum(params$groups))
    stop("You must specify groups in 'formula' for all the columns in 'data'")
  
  # Call FactoMineR::MFA() with corresponding arguments
  FactoMineR::MFA(base = data, group = params$groups, type = params$types,
    ind.sup = suprow, ncp = nd, name.group = params$names,
    num.group.sup = params$suppl, ..., graph = graph)
}

# x is a list with: expr, groups, types, names, suppl
# Note: this is a mess, I need to clean up this code!
get_groups <- function(x) {
  items <- as.list(x$expr)
  # We should have here '+', '-' or something else in case of last expression
  if (items[[1]] == "+") {
    x$suppl <- c(FALSE, x$suppl) # Not a supplementary variable
    x$expr <- items[[2]] # Second item is next expression
    # Third item is the information for that group
    # (n, or n * type, or n * type %as% name)
    item <- as.list(items[[3]])
  } else if (items[[1]] == "-") {
    x$suppl <- c(TRUE, x$suppl) # A supplementary variable
    x$expr <- items[[2]] # Second item is next expression
    # Third item is the information for that group
    # (n, or n * type, or n * type %as% name)
    item <- as.list(items[[3]])
  } else {
    # Last expression
    x$expr <- NULL
    item <- items
  }
  if (length(item) < 3) { # either n, or +n or -n
    if (length(item) == 2) {
      x$suppl <- switch(as.character(item[[1]]),
        "-" = c(TRUE, x$suppl),
        "+" = c(FALSE, x$suppl),
        stop("Bad 'formula', see help"))
      n <- item[[2]]
    } else n <- item[[1]]
    if (!is.numeric(n)) 
      stop("Bad 'formula', see help")
    x$groups <- c(as.integer(n, x$groups))
    x$types <- c("s", x$types)
    x$names <- c(NA, x$names)
  } else if (item[[1]] == "*") { # n * type or n * type %as% name
    if (length(item[[2]]) > 2)
      stop("Bad 'formula', see help")
    if (length(item[[2]]) == 2) {
      x$suppl <- switch(as.character(item[[2]][[1]]),
        "-" = c(TRUE, x$suppl),
        "+" = c(FALSE, x$suppl),
        stop("Bad 'formula', see help"))
      n <- item[[2]][[2]]
    } else n <- item[[2]]
    if (!is.numeric(n))
      stop("Bad 'formula', see help")
    x$groups <- c(as.integer(n), x$groups)
    # Right-hand side after '*': type or type %as% name
    subitem <- as.list(item[[3]])
    if (subitem[[1]] == "%as%") { # type %as name
      x$types <- c(as.character(subitem[[2]]), x$types)
      x$names <- c(as.character(subitem[[3]]), x$names)
    } else if (length(subitem) == 1 && is.name(subitem[[1]])) { # type only
      x$types <- c(as.character(subitem[[1]]), x$types)
      x$names <- c(NA, x$names)
    } else stop("Bad 'formula', see help") # Error
  } else if (item[[1]] == "%as%") { # n %as% name
    if (length(item[[2]]) == 2) {
      x$suppl <- switch(as.character(item[[2]][[1]]),
        "-" = c(TRUE, x$suppl),
        "+" = c(FALSE, x$suppl),
        stop("Bad 'formula', see help"))
      n <- item[[2]][[2]]
    } else n <- item[[2]]
    if (!is.numeric(n))
      stop("Bad 'formula', see help")
    x$groups <- c(as.integer(n), x$groups)
    x$types <- c("s", x$types) # type by default (s)
    x$names <- c(as.character(item[[3]]), x$names)
  } else stop("Bad 'formula', see help") # Error
  # Is there another expression to evaluate?
  if (!is.null(x$expr)) {
    x <- get_groups(x)
  } else {
    # Get correct vector for suppl
    if (length(x$suppl) < length(x$groups))
      x$suppl <- c(FALSE, x$suppl)
    # We need the group indices for the supplementary variables instead
    x$suppl <- (1:length(x$groups))[x$suppl]
    if (!length(x$suppl)) x$suppl <- NULL
    # Fix names
    def_names <- paste("group", 1:length(x$groups), sep = ".")
    names <- x$names
    names[is.na(names)] <- def_names[is.na(names)]
    x$names <- names
  }
  x
}
#form <- ~ -4 %as% Environnement -24*f %as% Plancton +1*f +18
#get_groups(list(expr = rlang::f_rhs(form)))

# TODO: methods for broom...

autoplot.MFA <- function(object,
type = c("screeplot", "altscreeplot", "loadings", "scores", "groups", "axes", "frequences", "ellipses"),
choices = 1L:2L, name = deparse(substitute(object)), col = "black",
fill = "gray", title, ..., env) {
  type = match.arg(type)

  if (missing(title))
    title <- paste(name, type, sep = " - ")

  res <- switch(type,
    screeplot = object %>.% # Classical screeplot
      tibble::tibble(eig = .$eig[, 1], PC = 1:nrow(.$eig)) %>.%
      chart(data = ., eig ~ PC) +
      geom_col(col = col, fill = fill) +
      labs(y = "Eigenvalues", title = title),

    altscreeplot = object %>.% # screeplot represented by dots and lines
      tibble::tibble(eig = .$eig[, 1], PC = 1:nrow(.$eig)) %>.%
      chart(data = ., eig ~ PC) +
      geom_line(col = col) +
      geom_point(col = "white", fill = col, size = 2, shape = 21, stroke = 3) +
      labs(y = "Eigenvalues", title = title),

    loadings = object %>.% # Plots of the variables
      plot(., axes = choices, choix = "var", title = title, ...),

    scores = object %>.% # Plot of the individuals
      plot(., axes = choices, choix = "ind", title = title, ...),

    groups = object %>.% # Plot of the groups
      plot(., axes = choices, choix = "group", title = title, ...),
    
    axes = object %>.% # Plot of the loadings for the various groups
      plot(., axes = choices, choix = "axes", title = title, ...),
    
    frequences = object %>.% # Plot of the correpondance analyses
      plot(., axes = choices, choix = "freq", title = title, ...),
    
    ellipses = object %>.% # Plot of the individuals with ellipses
      FactoMineR::plotellipses(., axes = choices, title = title, ...),

    stop("Unrecognized type, must be 'screeplot', 'altscreeplot',
      'loadings', 'scores', 'groups', 'axes', 'frequences', or 'ellipses'")
  )
  if (inherits(res, "ggplot")) {
    res
  } else {
    invisible(res)
  }
}

chart.MFA <- function(data, choices = 1L:2L, name = deparse(substitute(data)),
..., type = NULL, env = parent.frame())
  autoplot.MFA(data, choices = choices, name = name, ..., type = type, env = env)
class(chart.MFA) <- c("function", "subsettable_type")

# A hack to get fun$type() working in learnr
chart <- list(
  screeplot = function(data, type, ...) chart(data, type = "screeplot", ...),
  loadings = function(data, type, ...) chart(data, type = "loadings", ...),
  scores = function(data, type, ...) chart(data, type = "scores", ...),
  axes = function(data, type,...) chart(data, type = "axes",...)
)
```

```{r, echo=FALSE}
BioDataScience2::learnr_banner()
```

```{r, context="server"}
BioDataScience2::learnr_server(input, output, session)
```

------------------------------------------------------------------------

## Objectif

L'Analyse Factorielle Multiple (AFM) est une méthode statistique exploratoire qui permet de travailler avec plusieurs tableaux de données. Cette méthode permet également de faire le lien entre l'analyse en composantes principales (ACP) et l'Analyse factorielle des correspondances (AFC). La première méthode permet de travailler uniquement sur des variables quantitatives et la seconde uniquement sur des variables qualitatives. L'AFM est capable de traiter les deux types de variables simultanément. Cette méthode est par exemple employée en écologie. Sur une station un ensemble de mesures physico-chimiques peuvent être réalisée ainsi qu'un relevé de la faune et/ou de la flore.

Il est donc utile de comprendre la logique et l'usage de cette analyse. Ce tutoriel vous permettra de :

-   Réaliser de manière guidée une AFM

-   Effectuer les graphiques associées à cette analyse

-   Vous préparer à interpréter par vous-même les résultats de vos AFM

Avant toute chose, assurez vous d'avoir bien compris le contenu de la [section 8.1](https://wp.sciviews.org/sdd-umons2/?iframe=wp.sciviews.org/sdd-umons2-2020/analyse-factorielle-multiple-afm.html) du cours de SDD II.

## Étude de composition floristique

Des chercheurs ont étudié la composition floristique et les variables environnementales dans des forêts de pins sylvestres (*Pinus sylvestris*) dans l'est de la Fennoscandie, une région située au nord de l'Europe constituée de la Finlande, de la péninsule Scandinave, de la Carélie et de la péninsule de Kola.

![Fénoscandie, by [Mapsed](https://fr.wiktionary.org/wiki/Fennoscandie#/media/Fichier:Fennoscandia.jpg), license CC-BY-SA 4.0](images/Fennoscandia.jpg)

Les 24 sites sélectionnés pour cette étude sont situés dans le nord de la Finlande et dans la péninsule de Kola. Quatorze variables environnementales ont été mesurées sur les différents sites de l'étude, dont onze permettant d'évaluer la quantité de nutriments présentant dans la matière organique (MO). Les différentes variables mesurées sont présentées dans le tableau ci-dessous.

|    label     | description                    | unités     |
|:------------:|:-------------------------------|:-----------|
|    **N**     | azote                          | mg/g de MO |
|    **P**     | phosphore                      | µg/g de MO |
|    **K**     | potassium                      | µg/g de MO |
|    **Ca**    | calcium                        | µg/g de MO |
|    **Mg**    | magnésium                      | µg/g de MO |
|    **S**     | soufre                         | µg/g de MO |
|    **Al**    | aluminium                      | µg/g de MO |
|    **Fe**    | fer                            | µg/g de MO |
|    **Mn**    | manganèse                      | µg/g de MO |
|    **Zn**    | zinc                           | µg/g de MO |
|    **Mo**    | molybdène                      | µg/g de MO |
| **Baresoil** | surface estimée de sol nu      | \%         |
| **Humdepth** | épaisseur de la couche d'humus | cm         |
|    **pH**    | pH du sol                      | \-         |

Ces données sont disponible depuis `varechem` dans le package {vegan} et nous en réaliserons un tableau nommé `envir`.

```{r, echo=TRUE}
read("varechem", package = "vegan") %>.%
  rename(., station = .rownames) %->%
  envir
head(envir)
```

Les scientifiques ont également réalisé un relevé floristique (en excluant *Pinus sylvestris* du relevé). On y observe une couverture en pourcentage pour chaque station de chacune des 44 espèces étudiées. Ces données sont également disponible depuis `varespec` dans le package {vegan} que nous lirons dans la variable `spec`.

```{r, echo=TRUE}
read("varespec", package = "vegan") %>.%
  rename(., station = .rownames) %->%
  spec
head(spec)
```

Les données mises à disposition proviennent de l'étude suivante [Väre, H., Ohtonen, R. and Oksanen, J. (1995) Effects of reindeer grazing on understorey vegetation in dry Pinus sylvestris forests. Journal of Vegetation Science 6, 523--530](https://www.researchgate.net/publication/227830523_Effects_of_reindeer_grazing_on_vegetation_in_dry_Pinus_sylvestris_forests).

## Préparation du tableau de données

La première étape consiste dans la réalisation d'un tableau de données unique. Vous allez donc combiner le tableau `envir` avec le tableau `spec` en un seul objet nommé `vare` en spécifiant la colonne similaire entre ces deux tableaux qui est la variable `station` et qui servira de point de référence à la fusion. Affichez ensuite les premières lignes du tableau `vare` avec la fonction `head()`.

```{r join}
read("varechem", package = "vegan") %>.%
  rename(., station = .rownames) %->%
  envir
read("varespec", package = "vegan") %>.%
  rename(., station = .rownames) %->%
  spec
```

```{r join_h2, exercise=TRUE, exercise.setup="join"}
vare %<-% left_join(___ , ___, by = "___")
head(___)
```

```{r join_h2-hint-1}
vare %<-% left_join(envir , spec, by = "___")
head(___)

#### ATTENTION: Hint suivant = solution !####
```

```{r join_h2-solution}
## Solution ##
vare %<-% left_join(envir , spec, by = "station")
head(vare)
```

```{r join_h2-check}
grade_code("Maintenant que vous avez assemblé vos deux tableaux, vous allez pouvoir réaliser votre MFA. Il reste toutefois encore une petite étape avant cela.")
```

Il va encore falloir convertir la variable `station` en nom de ligne grâce à la fonction `column_to_rownames()`. Les numéros de stations seront utilisés par la suite sur la représentation dans l'espace de vos stations.

```{r rowname}
read("varechem", package = "vegan") %>.%
  rename(., station = .rownames) %->%
  envir
read("varespec", package = "vegan") %>.%
  rename(., station = .rownames) %->%
  spec

vare %<-% left_join(envir , spec, by = "station")
```

```{r rowname_h2, exercise=TRUE, exercise.setup="rowname"}
vare <- column_to_rownames(___, var = "___")
# Vérification des noms des lignes avec rownames()
rownames(___)
```

```{r rowname_h2-hint-1}
vare <- column_to_rownames(vare, var = "___")
# Vérification des noms des lignes avec rownames()
rownames(vare)

#### ATTENTION: Hint suivant = solution !####
```

```{r rowname_h2-solution}
## Solution ##
vare <- column_to_rownames(vare, var = "station")
# Vérification des noms des lignes avec rownames()
rownames(vare)
```

```{r rowname_h2-check}
grade_code("Bravo ! Vous venez de convertir la colonne station en nom de lignes. Cette astuce peut être utile dans certain cas. Il y a également la fonction opposée `rownames_to_column()`. Tout est prêt pour réaliser l'AFM à présent.")
```

## Réalisation de la AFM

L'AFM peut s'obtenir grâce à la fonction `mfa()`. Cette dernière requiert l'utilisation d'une formule avec une syntaxe un peu particulière. Votre formule commence par un `~` puis elle est composée de blocs qui doivent contenir trois éléments :

1.  Le nombre **n** de colonnes constituant le groupe
2.  Le type des variables qui composent le groupe. Vous avez quatre possibilités :
    -   **s** : Ce sont des variables quantitatives que l'on souhaite standardiser et les traiter par une ACP.

    -   **c** : Ce sont des variables quantitatives continues que l'on ne souhaite pas standardiser et les traiter par une ACP.

    -   **f** : Ce sont des variables qualitatives facteurs associées à des tableaux de contingences qu'on traite par une AFC.

    -   **n** : Ce sont des variables qualitatives de type dénombrements qu'on traite par une AFC.
3.  Le nom à donner au groupe

Le tableau `vare` est composé de quatorze variables associées à des mesures physico-chimiques et 44 variables liées à des relevés floristiques. Votre formule doit comprendre les éléments suivants pour que votre analyse soit la plus proche possible de celle que les auteurs ont réalisée dans la publication :

-   Le premier bloc doit se nommer **physico**, être traité par une ACP et les variables doivent être standardisées.
-   Le second bloc doit se nommer **flore**, être traité par une ACP et les variables ne doivent pas être standardisées.

Réalisez l'AFM demandée en utilisant le tableau `vare` puis affichez un résume de l'objet `vare_mfa` obtenu avec la fonction `summary()`

```{r varemfa}
read("varechem", package = "vegan") %>.%
  rename(., station = .rownames) %->%
  envir
read("varespec", package = "vegan") %>.%
  rename(., station = .rownames) %->%
  spec
vare %<-% left_join(envir , spec, by = "station")
vare <- column_to_rownames(vare, var = "station")
```

```{r mfa_h3, exercise=TRUE, exercise.setup="varemfa"}
vare_mfa <- mfa(data = ___, ~ ___*___ %as% ___ + ___*___ %as% ___)
summary(___)
```

```{r mfa_h3-hint-1}
vare_mfa <- mfa(data = vare, ~ ___*___ %as% ___ + ___*___ %as% ___)
summary(vare_mfa)
```

```{r mfa_h3-hint-2}
vare_mfa <- mfa(data = vare, ~ 14*s %as% physico + ___*___ %as% ___)
summary(vare_mfa)

#### ATTENTION: Hint suivant = solution !####
```

```{r mfa_h3-solution}
## Solution ##
vare_mfa <- mfa(data = vare, ~ 14*s %as% physico + 44*c %as% flore)
summary(vare_mfa)
```

```{r mfa_h3-check}
grade_code("Le résumé de votre analyse est complexe. Intéressez-vous tout particulièrement à la part de variance de chaque axe. Ces informations vous seront utile pour répondre à la question suivante.")
```

```{r variance_quiz}
question("Quelle est la proportion de la variance cumulée sur les deux premiers axes de l'AFM ?",
  answer("58%", correct = TRUE),
  answer("23%"),
  answer("34%"), 
  correct = "Bien joué ! Les deux premiers axes de l'AFM expriment 58% de la variance. Gardez à l'esprit que cette valeur est assez faible. Il serait aussi intéressant de considérer les trois premiers axes qui expriment 68% de la variance totale.",
  incorrect = "Oups, ce n'est pas la bonne réponse. Afin de trouver la bonne réponse, intéressez-vous au tableau `Eigenvalues`",
  allow_retry = TRUE
    )
```

### Graphique des éboulis

Réalisez un graphique des éboulis sur l'objet `vare_mfa` que vous avez réalisé précédemment.

```{r vare_mfa}
read("varechem", package = "vegan") %>.%
  rename(., station = .rownames) %->%
  envir
read("varespec", package = "vegan") %>.%
  rename(., station = .rownames) %->%
  spec

vare %<-% left_join(envir , spec, by = "station")
vare <- column_to_rownames(vare, var = "station")

vare_mfa <- mfa(data = vare, ~ 14*s %as% physico + 44*c %as% flore)
```

```{r scree_h2, exercise=TRUE, exercise.setup="vare_mfa", exercise.lines=1}
chart$___(___)
```

```{r scree_h2-hint-1}
chart$___(vare_mfa)

#### ATTENTION: Hint suivant = solution !####
```

```{r scree_h2-solution}
## Solution ##
chart$scree(vare_mfa)
```

```{r scree_h2-check}
grade_code("Ce graphique des éboulis permet de voir la part de variance exprimée sur chaque composante principale. Nous voyons bien que le premier axe reprend une très grande part de variance. On observe après un saut p pour le troisième axe est une diminution progressive de la variance exprimée sur les axes suivants.")
```

### Représentation des variables

Réalisez un graphique permettant de visualiser les variables dans le premier plan de l'AFM.

```{r loadings_h2, exercise=TRUE, exercise.setup="vare_mfa", exercise.lines=1}
chart$___(___ , choices = c(___, ___))
```

```{r loadings_h2-hint-1}
chart$___(vare_mfa, choices = c(___, ___))

#### ATTENTION: Hint suivant = solution !####
```

```{r loadings_h2-solution}
## Solution ##
chart$loadings(vare_mfa, choices = c(1, 2))
```

```{r loadings_h2-check}
grade_code("Ce graphique est assez chargé car il combine les vecteurs du groupe \"physico\" et du groupe \"flore\". Il s'analyse comme pour une ACP via la visualisation de l'importance des variables intiales dans le plan de l'AFM sous forme de vecteurs. Il est indispensable pour interpréter le graphique suivant qui répartit les observations dans le même plan. La norme (longueur) du vecteur indique si la variable est bien représentée dans ce plan ou non. Plus la norme du vecteur se rapproche de un, matérialisé par le cercle, mieux c'est. Celles aux normes trop faibles ne sont pas considérées dans l'analyse, pour ce plan-là en tous cas. Ensuite, les variables qui pointent dans la même direction sont corrélées positivement. Les variables qui pointent dans le sens opposé sont inversément corrélées. Les vecteurs orthogonaux correspondent à des variables non ou très faiblement corrélées entre elles.")
```

### Représentation des plans des ACP

L'AFM est un compromis réalisé entre les analyses indépendantes des différents tableaux. Réalisez une représentation des plans des ACP.

```{r axes_h2, exercise=TRUE, exercise.setup="vare_mfa", exercise.lines=1}
chart$___(___, choices = c(___, ___))
```

```{r axes_h2-hint-1}
chart$___(vare_mfa, choices = c(___, ___))

#### ATTENTION: Hint suivant = solution !####
```

```{r axes_h2-solution}
## Solution ##
chart$axes(vare_mfa, choices = c(1, 2))
```

```{r axes_h2-check}
grade_code("Ce graphique est similaire au précédent. Il s'interprète de la même façon. Cependant, on ne s'intéresse plus à la projection de chaque variable initiale mais des plans des ACP pour chaque groupe de variables. La première dimension du groupe \"flore\" pointe dans la même direction que la seconde dimension du groupe \"physico\". Il y a donc un forte corrélation entre ces deux plans. On obtient le même résultat pour le premier axe de \"physico\" et du second axe de \"flore\". Les normes des vecteurs sont proche de 1 pour chacune des ACP. Les deux ACP sont représentées correctement dans l'AFM.")
```

### Représentation des individus

Réalisez un graphique permettant de visualiser les individus dans le premier plan de l'AFM constitué des deux premières composantes principales.

```{r scores_h2, exercise=TRUE, exercise.setup="vare_mfa", exercise.lines=1}
chart$___(___, choices = c(___, ___))
```

```{r scores_h2-hint-1}
chart$___(vare_mfa, choices = c(___, ___)) 

#### ATTENTION: Hint suivant = solution !####
```

```{r scores_h2-solution}
## Solution ##
chart$scores(vare_mfa, choices = c(1, 2))
```

```{r scores_h2-check}
grade_code("La forme du nuage de points et surtout des clusters sont à rechercher ici. Nous pouvons voir par exemple un regroupement des stations 5, 6 et 7. On retrouve également un groupe particulier composé des stations 2, 3, 4, 9, 10, 11 et 12. Attention, les numéros des stations ne sont pas ici une indication de la proximité des stations.")
```

## Interprétation de l'AFM

La représentation des variables ainsi que la représentation de stations sont deux graphiques complémentaires afin d'interpréter l'AFM. Gardez à l'esprit que la part de variance exprimée par les deux premiers axes reste relativement faible avec 58%.

```{r}
read("varechem", package = "vegan") %>.%
  rename(., station = .rownames) %->%
  envir
read("varespec", package = "vegan") %>.%
  rename(., station = .rownames) %->%
  spec

vare %<-% left_join(envir , spec, by = "station")
vare <- column_to_rownames(vare, var = "station")

vare_mfa <- mfa(data = vare, ~ 14*s %as% physico + 44*c %as% flore)

chart$scores(vare_mfa, choices = c(1, 2))
```

```{r}
chart$loadings(vare_mfa, choices = c(1, 2))
```

Les stations 2, 3, 4, 9, 10, 11 et 12 sont des stations particulières qui ne sont pas des zones de pâturages. On retrouve dans ces stations une dominance de *Cladina stellaris* alors que cet espèce est très peu présente dans les zones de pâturages. Les zones de pâturages sont caractérisée par une haute concentration en azote. Les stations 5, 6 et 7 sont des zones de pâturages avec une forte teneur en azote et une présence massive de *Cladina arbuscala*.

Nous n'irons pas plus loin dans l'interprétation de cette AFM et vous laissons la possibilité de lire la publication liée à cette étude. Les scientifiques ont réalisé une NMDS en fait, pas une AFM, mais les résultats obtenus sont assez proches.

## Conclusion

Cette AFM vous a permis d'analyser deux tableaux contenant des variables quantitatives, traités différemment (l'un standardisé, l'autre pas) en même temps. L'AFM permet aussi d'y adjoindre des groupes qualitatifs avec des analyses de type AFC. Cette méthode multi-tableaux est donc très polyvalente (mais notez qu'il en existe d'autres aussi).

Maintenant que vous avez compris la logique et que vous êtes capable d'écrire le code permettant de réaliser une AFM, vous pouvez appliquer cette technique par vous-même (assignation GitHub dans le cours).

```{r comm_noscore, echo=FALSE}
question_text(
  "Laissez-nous vos impressions sur cet outil pédagogique",
  answer("", TRUE, message = "Pas de commentaires... C'est bien aussi."),
  incorrect = "Vos commentaires sont enregistrés.",
  placeholder = "Entrez vos commentaires ici...",
  allow_retry = TRUE
)
```
