---
title: "Modèle linéaire"
author: "Guyliann Engels & Philippe Grosjean"
description: "**SDD II Module 3** Application des concepts liés au modèle linéaire."
tutorial:
  id: "B03La_mod_lin"
  version: 2.1.0/8
output: 
  learnr::tutorial:
    progressive: true
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
BioDataScience2::learnr_setup()
SciViews::R()
```

```{r, echo=FALSE}
BioDataScience2::learnr_banner()
```

```{r, context="server"}
BioDataScience2::learnr_server(input, output, session)
```

----

## Objectifs

- Maîtriser le modèle linéaire et l'ANCOVA

## Modèle linéaire 1

```{r mais1_init}
set.seed(42)
x <- seq(from = 5, to = 25, by = 0.25)
a <- x*1 + 3 + rnorm(sd = 0.5, n = length(x))
b <- x*1.1 + 2.5 + rnorm(sd = 0.5, n = length(x))
c <- x*1.2 + 2 + rnorm(sd = 0.5, n = length(x))

area <- as.factor(rep(c("a", "b", "c"), each = length(x)))

mais1 <- tibble(
  x = c(x,x,x),
  value = c(a,b,c),
  area = area
)
```

Réalisez la régression linéaire multiple la plus adpatée de `value` en fonction de `x` et de la variable `area` sur le jeu de données `mais1`. Vous avez à votre dispositon un nuage de points et un résumé des données pour prendre connaissance des données.

*si tu souhaites faire passer ton modèle linéaire par le point 0,0, utilise la convention.*

```{r}
lm1 <- lm(data = mais1, value ~ x * area)
#summary(lm1)
lm1_param <- broom::glance(lm1)
lm1_result <- broom::tidy(lm1)

chart(mais1, value ~ x %col=%area) +
  geom_point()

summary(mais1)
```

```{r mais1_prep}
# coppy of mais1_init
set.seed(42)
x <- seq(from = 5, to = 25, by = 0.25)
a <- x*1 + 3 + rnorm(sd = 0.5, n = length(x))
b <- x*1.1 + 2.5 + rnorm(sd = 0.5, n = length(x))
c <- x*1.2 + 2 + rnorm(sd = 0.5, n = length(x))

area <- as.factor(rep(c("a", "b", "c"), each = length(x)))

mais1 <- tibble(
  x = c(x,x,x),
  value = c(a,b,c),
  area = area
)
```

```{r mais1_h3, exercise = TRUE, exercise.setup = "mais1_prep", exercise.lines = 1}
summary(lm. <- lm(data = ___, ___))
```

```{r mais1_h3-hint-1}
# chunk
summary(lm. <- lm(data = DF, FORMULA))
```

```{r mais1_h3-hint-2}
# chunk
summary(lm. <- lm(data = DF, VAR1 ~ VAR2*VAR3))

#### ATTENTION: Hint suivant = solution !####
```

```{r mais1_h3-solution}
summary(lm. <- lm(data = mais1, value ~ x * area))
```

```{r mais1_h3-check}
grade_code("Félicitation ! Tu viens de réaliser ton premier modèle linéaire.")
```

```{r qu_mais1}
question("Quelle est la valeur de la pente du modèle linéaire du niveau `a` de la variable `area` ?",
         answer(sprintf("%.3f", lm1_result$estimate[[2]]), correct = TRUE),
           answer(sprintf("%.3f", lm1_param$AIC)),
           answer(sprintf("%.3f", lm1_result$estimate[[3]])),
           answer(sprintf("%.3f", lm1_result$estimate[[1]])),
           allow_retry = TRUE, random_answer_order = TRUE
           )
```

## Modèle linéaire 2

```{r mais2_init}
set.seed(42)
x <- seq(from = 5, to = 25, by = 0.25)
a <- x*1 + 3 + rnorm(sd = 0.5, n = length(x))
b <- x*1.1 + 3 + rnorm(sd = 0.5, n = length(x))
c <- x*1.2 + 3 + rnorm(sd = 0.5, n = length(x))

area <- as.factor(rep(c("a", "b", "c"), each = length(x)))

mais2 <- tibble(
  x = c(x,x,x),
  value = c(a,b,c),
  area = area
)
```

Réalisez la régression linéaire multiple la plus adpatée de `value` en fonction de `x` et de la variable `area` sur le jeu de données `mais2`. Vous avez à votre dispositon un nuage de points et un résumé des données pour prendre connaissance des données.

*si tu souhaites faire passer ton modèle linéaire par le point 0,0, utilise la convention.*

```{r}
chart(mais2, value ~ x %col=% area) +
  geom_point()

summary(mais2)

#lm2 <- lm(data = mais2, value ~ x * area) # modèle le plus compliqué doit être simplifié 
#lm2 <- lm(data = mais2, value ~ x + area + x:area)
lm2 <- lm(data = mais2, value ~ x + x:area)
#summary(lm2)

lm2_param <- broom::glance(lm2)
lm2_result <- broom::tidy(lm2)
```

```{r mais2_prep}
set.seed(42)
x <- seq(from = 5, to = 25, by = 0.25)
a <- x*1 + 3 + rnorm(sd = 0.5, n = length(x))
b <- x*1.1 + 3 + rnorm(sd = 0.5, n = length(x))
c <- x*1.2 + 3 + rnorm(sd = 0.5, n = length(x))

area <- as.factor(rep(c("a", "b", "c"), each = length(x)))

mais2 <- tibble(
  x = c(x,x,x),
  value = c(a,b,c),
  area = area
)
```

```{r mais2_h4, exercise = TRUE, exercise.setup = "mais2_prep", exercise.lines = 1}
summary(lm. <- lm(data = ___, ___))
```

```{r mais2_h4-hint-1}
# chunk
summary(lm. <- lm(data = DF, FORMULA))
```

```{r mais2_h4-hint-2}
# chunk
summary(lm. <- lm(data = DF, VAR1 ~ VAR2*VAR3))
# ou encore
summary(lm. <- lm(data = DF, VAR1 ~ VAR2 + VAR3 + VAR2:VAR3))
```

```{r mais2_h4-hint-3}
summary(lm. <- lm(data = mais2, value ~ x * area))
# ou 
summary(lm. <- lm(data = mais2, value ~ x + area + x:area))

#### ATTENTION: Hint suivant = solution !####
```

```{r mais2_h4-solution}
summary(lm. <- lm(data = mais2, value ~ x + x:area))
```

```{r mais2_h4-check}
grade_code("Tu progresses bien ! Plus que quelques petits notions à apprendre pour bien maitriser la matière.")
```

```{r qu_mais2}
question("Quelle est la valeur de la pente du modèle linéaire du niveau `c` de la variable `area` ?",
         answer(sprintf("%.3f", (lm2_result$estimate[[2]] + lm2_result$estimate[[4]])), correct = TRUE),
           answer(sprintf("%.3f", lm2_param$AIC)),
           answer(sprintf("%.3f", lm2_result$estimate[[4]])),
           answer(sprintf("%.3f", lm2_result$estimate[[1]])),
           allow_retry = TRUE, random_answer_order = TRUE
           )
```

## Modèle linéaire 3

```{r mais3_init}
set.seed(42)
x <- seq(from = 5, to = 25, by = 0.25)
a <- x*1 + 2 + rnorm(sd = 0.5, n = length(x))
b <- x*1 + 6 + rnorm(sd = 0.5, n = length(x))
c <- x*1 + 4 + rnorm(sd = 0.5, n = length(x))

area <- as.factor(rep(c("a", "b", "c"), each = length(x)))

mais3 <- tibble(
  x = c(x,x,x),
  value = c(a,b,c),
  area = area
)
```

Réalisez la régression linéaire multiple la plus adpatée de `value` en fonction de `x` et de la variable `area` sur le jeu de données `mais3 `. Vous avez à votre dispositon un nuage de points et un résumé des données pour prendre connaissance des données.

*si tu souhaites faire passer ton modèle linéaire par le point 0,0, utilise la convention `-1`*

```{r}
#lm3 <- lm(data = mais3, value ~ x*area)
#lm3 <- lm(data = mais3, value ~ x + area + x:area)
lm3 <- lm(data = mais3, value ~ x + area)
#summary(lm3)

lm3_param <- broom::glance(lm3)
lm3_result <- broom::tidy(lm3)

chart(mais3, value ~ x %col=%area) +
  geom_point()

summary(mais3)
```

```{r mais3_prep}
# copy of mais3_init
set.seed(42)
x <- seq(from = 5, to = 25, by = 0.25)
a <- x*1 + 2 + rnorm(sd = 0.5, n = length(x))
b <- x*1 + 6 + rnorm(sd = 0.5, n = length(x))
c <- x*1 + 4 + rnorm(sd = 0.5, n = length(x))

area <- as.factor(rep(c("a", "b", "c"), each = length(x)))

mais3 <- tibble(
  x = c(x,x,x),
  value = c(a,b,c),
  area = area
)
```

```{r mais3_h4, exercise = TRUE, exercise.setup = "mais3_prep", exercise.lines = 1}
summary(lm. <- lm(data = ___, ___))
```

```{r mais3_h4-hint-1}
# chunk
summary(lm. <- lm(data = DF, FORMULA))
```

```{r mais3_h4-hint-2}
# chunk
summary(lm. <- lm(data = DF, VAR1 ~ VAR2*VAR3))
# ou encore
summary(lm. <- lm(data = DF, VAR1 ~ VAR2 + VAR3 + VAR2:VAR3))
```

```{r mais3_h4-hint-3}
summary(lm. <- lm(data = mais2, value ~ x * area))
# ou 
summary(lm. <- lm(data = mais2, value ~ x + area + x:area))

#### ATTENTION: Hint suivant = solution !####
```

```{r mais3_h4-solution}
summary(lm. <- lm(data = mais3, value ~ x + area))
```

```{r mais3_h4-check}
grade_code("Continue comme ca. C'est excellent !")
```

```{r qu_mais3}
question("Quelle est la valeur de l'ordonnée à l'origine du modèle linéaire du niveau `b` de la variable `area` ?",
         answer(sprintf("%.3f", (lm3_result$estimate[[1]] + lm3_result$estimate[[3]])), correct = TRUE),
           answer(sprintf("%.3f", lm3_param$AIC)),
           answer(sprintf("%.3f", lm3_result$estimate[[4]])),
           answer(sprintf("%.3f", lm3_result$estimate[[1]])),
           allow_retry = TRUE, random_answer_order = TRUE
           )
```

## Modèle linéaire 4

```{r mais4_init}
set.seed(42)
x <- seq(from = 5, to = 25, by = 0.25)
a <- x*1 + rnorm(sd = 0.5, n = length(x))
b <- x*1.1  + rnorm(sd = 0.5, n = length(x))
c <- x*1.2  + rnorm(sd = 0.5, n = length(x))

area <- as.factor(rep(c("a", "b", "c"), each = length(x)))

mais4 <- tibble(
  x = c(x,x,x),
  value = c(a,b,c),
  area = area
)
```

Réalisez la régression linéaire multiple la plus adpatée de `value` en fonction de `x`  et de la variable `area` sur le jeu de données `mais4`. Vous avez à votre dispositon un nuage de points et un résumé des données pour prendre connaissance des données.

*si tu souhaites faire passer ton modèle linéaire par le point 0,0, utilise la convention `-1`*


```{r}
#lm4 <- lm(data = mais4, value ~ x * area)
#lm4 <- lm(data = mais4, value ~ x * area + 0)
lm4 <- lm(data = mais4, value ~ x + x:area - 1)
#
#summary(lm4)
lm4_param <- broom::glance(lm4)
lm4_result <- broom::tidy(lm4)

chart(mais4, value ~ x %col=%area) +
  geom_point()

summary(mais4)
```

```{r mais4_prep}
# copy of mais4_init
set.seed(42)
x <- seq(from = 5, to = 25, by = 0.25)
a <- x*1 + rnorm(sd = 0.5, n = length(x))
b <- x*1.1  + rnorm(sd = 0.5, n = length(x))
c <- x*1.2  + rnorm(sd = 0.5, n = length(x))

area <- as.factor(rep(c("a", "b", "c"), each = length(x)))

mais4 <- tibble(
  x = c(x,x,x),
  value = c(a,b,c),
  area = area
)
```

```{r mais4_h4, exercise = TRUE, exercise.setup = "mais4_prep", exercise.lines = 1}
summary(lm. <- lm(data = ___, ___))
```

```{r mais4_h4-hint-1}
# chunk
summary(lm. <- lm(data = DF, FORMULA))
```

```{r mais4_h4-hint-2}
# chunk
summary(lm. <- lm(data = DF, VAR1 ~ VAR2*VAR3))
# ou encore
summary(lm. <- lm(data = DF, VAR1 ~ VAR2 + VAR3 + VAR2:VAR3))
```

```{r mais4_h4-hint-3}
summary(lm. <- lm(data = mais2, value ~ x * area))
# ou 
summary(lm. <- lm(data = mais2, value ~ x + area + x:area))

#### ATTENTION: Hint suivant = solution !####
```

```{r mais4_h4-solution}
summary(lm. <- lm(data = mais4, value ~ x + x:area - 1))
```

```{r mais4_h4-check}
grade_code("Tu maitrises vraiment bien la modélisation linéaire !")
```

```{r qu_mais4}
question("Quelle est la valeur de l'ordonnée à l'origine du modèle linéaire du niveau `b` de la variable `area` ?",
         answer(sprintf("%.3f", 0), correct = TRUE),
           answer(sprintf("%.3f", lm4_param$AIC)),
           answer(sprintf("%.3f", lm4_result$estimate[[3]])),
           answer(sprintf("%.3f", lm4_result$estimate[[2]])),
           allow_retry = TRUE, random_answer_order = TRUE
           )
```

## Conclusion

Vous venez de terminer votre séance d'exercice.

```{r comm_noscore, echo=FALSE}
question_text(
  "Laissez-nous vos impressions sur ce learnr",
  answer("", TRUE, message = "Pas de commentaires... C'est bien aussi."),
  incorrect = "Vos commentaires sont enregistrés.",
  placeholder = "Entrez vos commentaires ici...",
  allow_retry = TRUE
)
```
